- interval: 10s
  input_series:
    - series: 'maas_net_subnet_ip_count{status="available", cidr="10.0.0.0/24"}'
      values: '255 254 160 90 10 5 0 0'
    - series: 'maas_net_subnet_ip_count{status="available", cidr="10.0.1.0/24"}'
      values: '255 254 160 90 10 15 80 10'
  alert_rule_test:
    - eval_time: 5m
      alertname: MAASIPExhaustionCritical
      exp_alerts:
        - exp_labels:
            severity: critical
            cidr: "10.0.0.0/24"
            status: "available"
          exp_annotations:
            summary: "10.0.0.0/24 has exhausted all IPs"
            description: "The subnet has no more IPs to allocate"
