- interval: 10s
  input_series:
    - series: 'maas_virsh_fetch_description_failure'
      values: '0 0 1 0 0 0 1 0 0 1 1'
  alert_rule_test:
    - eval_time: 1m
      alertname: MAASVirshMachineDescriptionFetchFailure
      exp_alerts:
        - exp_labels:
            severity: error
          exp_annotations:
            summary: "MAAS has failed to fetch a machine's description from Virsh"
            description: "MAAS has failed to communicate with a Virsh hypervisor while fetching a machine's description"
